# 11. More List Manipulation

You can do useful things with lists like adding them together or making lists of lists. Let's start with an example. The cell below declares two lists

```python
list_a = [1, 2, 3]
list_b = [4, 5, 6]
```

**Run the code blocks below and write what the operations do in the cell at the bottom.** 

(N.B. you may get some odd results if you run them out of order, if you do, rerun the cells starting from the cell above to re-declare list_a and list_b)

```python
# 1
list_a + list_b
```

```python
# 2
new_list = [list_a, list_b]
new_list
```

```python
# 3
list_a.append(list_b)
list_a
```

**What is the difference between the three lists?**

Write your notes in the cell below as a comment:

```python
#
```

By appending `list_b` to `list_a`, we've changed it, see for yourself

```python
list_a
```

So we're going to redeclare it here (this is why you should be careful in coding and set new variables and use functions):

```python
list_a=[1,2,3]
```

We'll be doing plots in a later notebook, but the code below shows you how we might plot our data. `list_a` is the $x$ values and `list_b` the $y$ values. 

`matplotlib` is another module, like `numpy` that we can import and use, it contains lots of functions for plotting data. Here we use a function called `scatter` that plots a scatter plot of the input data.

```python
import matplotlib.pyplot as plt
plt.scatter(list_a, list_b)
```

## Getting elements of a list

The objects (numbers, strings, whatever) in a list are the elements of that list. As well as looping over all of them (as we have done so far); It can be very useful to access specific elements directly so you can then do things with them. 

This section will demonstrate that, but first we have to measure the length of our list. 

As you recall, to find out the **len**gth of a list we use `len(list_name)`

```python
len(list_a)
```

And we can *index* (retrieve) elements from of a list. See below, what do you think is going on?

```python
list_a[0]
```

```python
list_b[0]
```

```python
list_b[3]
```

The cell above should have given you an `IndexError`, this is an *exception* that has been *raised* due to the index being used (3) being too high for the list being indexed (as the indices go from 0 to 2). This is an example of how error messages can help you to debug your code. We will cover exceptions in more detail in a future week.

We can **index** the elements of a list. As python is zero indexed, we start from 0, so a list of length 3 has the indices: `[0, 1, 2]`, and we can pull out the first item by writing `my_list[0]`, the second item by writing `my_list[1]` etc. 

What do you think the next chunk of code might do?

```python
list_b[-1]
```

```python
list_b[-2]
```

```python
list_b[-3]
```

We can start indexed from the end of the list, and pull out the elements in reverse order. 

As the list is 3 elements long, what do you think will happen if you input `list_b[-4]`? Try it out in the cell below

```python

```

Lists can also be **sliced** by using the `:`, this effectively extracts certain elements to make a new list. For example, if you wanted to slice from the start of a list of elements to index 3 (i.e. 0, 1, 2, 3):

```python
element_list = ['H', 'He', 'Li', 'Be', 'B', 'C', 'N', 'O', 'F', 'Ne']
```

```python
element_list[:4]
```

Try inverting the `:` and `4`.

```python
#
```

What do you notice about the slice? We can also slice between values, write the code for `element_list[-5:-3]`, what do you notice?

```python
#
```

You may have noticed that the indices of a list are also a list. Lets see if we can generate them automatically (this might remind you of what was done in the enumerate section in the previous notebook)

```python
# list comprehension to generate a list of indices
[i for i in range(len(list_b))]
```

The code above is a **list comprehension** it does the same as a loop for simple lists, for example:

```python
# script to get indices of a list
index_list = []
for i in range(len(list_b)):
    index_list.append(i)
    
index_list
```

List comprehensions can be very useful for writing short code, but writing out the full loop does the same thing, so we'll primarily focus on that approach.

## **Worked problem**

Modify the code below to create a list of indices called `my_index_list` which contains the indices of `my_list`, and also creates a new list called `my_copied_list` which copies the elements of `my_list`. 

So, you should create:

`my_copied_list` which is [2,4,6,8,10]

`my_index_list` which is [0,1,2,3,4]

Don't forget to change which list you loop over and check the names of all the lists!

```python
my_list = [2, 4, 6, 8, 10]
```

```python
## modify this code
index_list = []
for i in range(len(list_b)):
    index_list.append(i)
```

The cells below will help you test your script.

```python
my_index_list
```

```python
my_copied_list
```

The cell below will allow me to test your code! This is an **assertion**, a statement of something that should be true that is used to test codes and raise errors.

<!-- grade_id: cell-d3e88426913218cb, points: 6, grade: true, solution: false -->
```python
assert my_copied_list == [2, 4, 6, 8, 10]
assert my_index_list == [0, 1, 2, 3, 4]
```

Now put your code in the function below. Just copy it in between the comment at the top and the return statement, don't forget to tab the code inwards and remove the NotImplemented Error.

<!-- grade_id: cell-ea23b9deb93733bb, points: 0, grade: false, solution: true -->
```python
def list_copier(my_list):
    """copies a list and returns the copy and a list of indices"""
    ## modify this code
    ###
    return my_copied_list, my_index_list
```

```python
my_copied_list, my_index_list = list_copier(my_list)
print(f"my_copied_list = {my_copied_list}")
print(f"my_index_list = {my_index_list}")
```

<!-- grade_id: cell-15c9456541f774d5, points: 4, grade: true, solution: false -->
```python
assert my_copied_list == [2, 4, 6, 8, 10]
assert my_index_list == [0, 1, 2, 3, 4]
my_copied_list, my_index_list = list_copier([1,2])
assert my_index_list == [0, 1]
assert my_copied_list == [1, 2]
```

## Making a list of pairs

You can put any type of python variable into a list, even other lists! This is very useful. 

The code below will make a list of pairs from the two lists we declared at the top of this document.

```python
list_of_pairs=[]
for i in range(len(list_a)):
    list_of_pairs.append([list_a[i], list_b[i]])
    
list_of_pairs
```

There's some new stuff in this example. 

We tell python that something is a list by putting it into square brackets, `[ ]`. For a list of lists we have more than one set of brackets, e.g.:

`[ [ first list ] , [ second list ], [ third list ] ]`

You can pull out elements of the list of lists exactly as you can pull out elements of a list. 

Try this out for yourself. In the cell below, try out

`list_of_pairs[x]`

and replace x by an index, (i.e. 0, -1, 3, etc)
and write what it does in your own words in the cell below.

```python
## write some code in this cell

#
```

For reference, here is the list comprehension that does the same thing as loop above:

```python
[[list_a[i], list_b[i]]for i in range(len(list_a))]
```

**You finished this notebook, now go on to number 12**