# 12. Dictionaries

The `list` data type is a great storage container in Python. Good... but limited.

Sometimes, we want to be able to store multiple items in one place, like a `list`, but we also want to create some **associations** inside the storage container. 

This is where the `dictionary` data type comes in.

A `dictionary` in Python associates a **key** with a **value**. This is like a *real* dictionary where a **word** is associated with a **definition**.

Hit `Shift+Enter` on the code block below to get started.

```python
# Define a function called periodictable that takes one input argument
def periodictable(elementnumber):
    
    # Define a dictionary to store atomic numbers as keys, and element symbols as the associated values
    ptable={
    'H': 1, 'He': 2, 'Li': 3, 'Be': 4, 'B': 5, 'C': 6, 'N': 7, 'O': 8, 'F': 9, 'Ne': 10,
    'Na': 11, 'Mg': 12, 'Al': 13, 'Si': 14, 'P': 15, 'S': 16, 'Cl': 17, 'Ar': 18,
    'K': 19, 'Ca': 20, 'Sc': 21, 'Ti': 22, 'V': 23, 'Cr': 24, 'Mn': 25, 'Fe': 26, 'Co': 27,
    'Ni': 28, 'Cu': 29, 'Zn': 30, 'Ga': 31, 'Ge': 32, 'As': 33, 'Se': 34, 'Br': 35, 'Kr': 36,
    'Rb': 37, 'Sr': 38, 'Y': 39, 'Zr': 40, 'Nb': 41, 'Mo': 42, 'Tc': 43, 'Ru': 44, 'Rh': 45,
    'Pd': 46, 'Ag': 47, 'Cd': 48, 'In': 49, 'Sn': 50, 'Sb': 51, 'Te': 52, 'I': 53, 'Xe': 54,
    'Cs': 55, 'Ba': 56, 'La': 57, 'Ce': 58, 'Pr': 59, 'Nd': 60, 'Pm': 61, 'Sm': 62, 'Eu': 63,
    'Gd': 64, 'Tb': 65, 'Dy': 66, 'Ho': 67, 'Er': 68, 'Tm': 69, 'Yb': 70, 'Lu': 71,
    'Hf': 72, 'Ta': 73, 'W': 74, 'Re': 75, 'Os': 76, 'Ir': 77, 'Pt': 78, 'Au': 79, 'Hg': 80,
    'Tl': 81, 'Pb': 82, 'Bi': 83, 'Po': 84, 'At': 85, 'Rn': 86, 'Fr': 87, 'Ra': 88,
    'Ac': 89, 'Th': 90, 'Pa': 91, 'U': 92, 'Np': 93, 'Pu': 94, 'Am': 95, 'Cm': 96, 'Bk': 97,
    'Cf': 98, 'Es': 99, 'Fm': 100, 'Md': 101, 'No': 102, 'Lr': 103, 'Rf': 104, 'Db': 105,
    'Sg': 106, 'Bh': 107, 'Hs': 108, 'Mt': 109, 'Ds': 110, 'Rg': 111, 'Cn': 112,
    'Nh': 113, 'Fl': 114
    }
   
    # Retrieve a value from the dictionary that is associated with the key given by elementnumber
    element=ptable[elementnumber]
    
    # Return the variable element for use elsewhere in the program
    return element

# Call the periodictable function
periodictable('H')
```

Now that you have run some code that uses the `dictionary` (or `dict`) data type, we can let you in on a secret that you may or may not have spotted already.

If you tried to call the `periodictable` function using the elements `Mc`, `Lv`, `Ts`, or `Og`, it would not work. Try it.

In other words, there are 4 elements missing from the above `dictionary` that represents the periodic table! They are:

* Moscovium (Mc): Element 115 
* Livermorium (Lv): Element 116 
* Tennessine (Ts): Element 117 
* Oganesson (Og): Element 118 

In the duplicated `dictionary` in the code block below, add the 4 missing elements, using the correct syntax for the `dictionary`.

Show that your updated code block works by calling the function, using one of the new elements.

<!-- grade_id: cell-ccea6efb98c22bfa, points: 0, grade: false, solution: true -->
```python
def periodictable(elementnumber):
    
    
    ptable={
    'H': 1, 'He': 2, 'Li': 3, 'Be': 4, 'B': 5, 'C': 6, 'N': 7, 'O': 8, 'F': 9, 'Ne': 10,
    'Na': 11, 'Mg': 12, 'Al': 13, 'Si': 14, 'P': 15, 'S': 16, 'Cl': 17, 'Ar': 18,
    'K': 19, 'Ca': 20, 'Sc': 21, 'Ti': 22, 'V': 23, 'Cr': 24, 'Mn': 25, 'Fe': 26, 'Co': 27,
    'Ni': 28, 'Cu': 29, 'Zn': 30, 'Ga': 31, 'Ge': 32, 'As': 33, 'Se': 34, 'Br': 35, 'Kr': 36,
    'Rb': 37, 'Sr': 38, 'Y': 39, 'Zr': 40, 'Nb': 41, 'Mo': 42, 'Tc': 43, 'Ru': 44, 'Rh': 45,
    'Pd': 46, 'Ag': 47, 'Cd': 48, 'In': 49, 'Sn': 50, 'Sb': 51, 'Te': 52, 'I': 53, 'Xe': 54,
    'Cs': 55, 'Ba': 56, 'La': 57, 'Ce': 58, 'Pr': 59, 'Nd': 60, 'Pm': 61, 'Sm': 62, 'Eu': 63,
    'Gd': 64, 'Tb': 65, 'Dy': 66, 'Ho': 67, 'Er': 68, 'Tm': 69, 'Yb': 70, 'Lu': 71,
    'Hf': 72, 'Ta': 73, 'W': 74, 'Re': 75, 'Os': 76, 'Ir': 77, 'Pt': 78, 'Au': 79, 'Hg': 80,
    'Tl': 81, 'Pb': 82, 'Bi': 83, 'Po': 84, 'At': 85, 'Rn': 86, 'Fr': 87, 'Ra': 88,
    'Ac': 89, 'Th': 90, 'Pa': 91, 'U': 92, 'Np': 93, 'Pu': 94, 'Am': 95, 'Cm': 96, 'Bk': 97,
    'Cf': 98, 'Es': 99, 'Fm': 100, 'Md': 101, 'No': 102, 'Lr': 103, 'Rf': 104, 'Db': 105,
    'Sg': 106, 'Bh': 107, 'Hs': 108, 'Mt': 109, 'Ds': 110, 'Rg': 111, 'Cn': 112,
    'Nh': 113, 'Fl': 114 # Added missing elements
    }
   
    
    element=ptable[elementnumber]
    
    return element
```

<!-- grade_id: cell-f2cf1ab7aea12493, points: 1, grade: true, solution: false -->
```python
assert periodictable('H') == 1
assert periodictable('Mc') == 115
```

<!-- grade_id: cell-2d890d8b6f504bae, points: 1, grade: true, solution: false -->
```python
assert periodictable('Ta') == 73
assert periodictable('Lv') == 116
```

<!-- grade_id: cell-453eff9c48de5f29, points: 1, grade: true, solution: false -->
```python
assert periodictable('He') == 2
assert periodictable('Ts') == 117
```

<!-- grade_id: cell-d4c2b84e0f950ead, points: 1, grade: true, solution: false -->
```python
assert periodictable('B') == 5
assert periodictable('Og') == 118
```

# A worked example: using dictionaries, lists and functions to calculate molecular weight easily

This final example will give you some idea of the power of the python you've been taught so far and show you how a coder goes about building up code. The aim of this example is to build a function that can calculate the molecular weight from a list of atoms in the molecule.

First, we shall initialise a dictionary called `weight_dict` to relate atomic symbols to the element's atomic weights:

```python
weight_dict = {'H': 1.008,
               'C': 12.011,
               'O': 15.999}
```

Now, we shall set up a dictionary called `mol_dict` to relate the name of a molecule to a list of the atoms in it. 

Note, there are many ways to input molecular formulae and this is not a good one! You'll learn about SMILES strings later on, but this will do for now.

```python
mol_dict = {'benzene': 'CHCHCHCHCHCH',
            'methanol': 'CHHHOH',
            'ethanol': 'CHHHCHHOH'}
```

Just as lists have the *empty list, []*, there is the *empty dictionary {}* - note that dictionaries use {} to distinguish them from lists.

Also, as per lists, you can assign values to keys directly. For instance, we could have written the previous code as follows:

```python
mol_dict = {}
mol_dict['benzene'] = 'CHCHCHCHCHCH'
mol_dict['methanol']= 'CHHHOH'
mol_dict['ethanol']= 'CHHHCHHOH'
```

When should you use one or the other? In general, if you can do the first style of 'big bang' initialisation, you should - most people would agree it is clearer to read. However, you may not be able to; e.g. if you are generating the values in a dictionary from a for loop iterating over a list.

### Other Dictionary Facts!
You can overwrite a value in a dictionary!
````
my_dict = {x:0, y:0}
my_dict[x]=1
# my_dict is now {x:1, y:0}
````
You can get a list* of the keys in a dictionary by calling `.keys()`, and a list* of the values in a dictionary using `.values()`

*Actually, like range, you get a generator, not a list, so you will need to treat it accordingly.

```python
mol_dict.keys()
```

```python
mol_dict.values()
```

It is also possible to extract a given value from a dictionary with the `get()` function, which follows the format `get(key, other)`

```python
mol_dict.get('benzene', 'None')
```

To develop with, we shall just pick one of our molecules to use to try out and we shall set the variable `mol_code` to that:

```python
mol_code=mol_dict['benzene']
mol_code
```

Now, a string of letters is also a list and can be addressed the way we address lists, i.e. to get the first letter of a string you can use `my_string[0]` and so on. Try it out 

```python
mol_code[0]
```

Part of the power of python is that you can **nest** code *statements*. 

For example, if I want the atomic weight of carbon, I would write:

`weight_dict['C']`

and if I wanted to first atom of benzene, I would write

`mol_code[0]`, which should return `'C'`

so we can nest these operations, like so:

`weight_dict[ mol_code[0] ]`

try it out in the cell below

```python

```

In fact, you don't need to stop there, if you wanted the mass of the 5th atom of methanol, you could write:
`mol_dict['methanol'][4]` to get the 'O' and put it into the construction above, i.e.


```python
weight_dict[ mol_dict['methanol'][4] ]
```

Now, we're going to write a little script to loop over the variable mol_code and print which atom is which:

```python
count = 0
while count < len(mol_code):
    atom = mol_code[count]
    print(f'atom {count} is a {atom}')
    count = count + 1
```

This little script uses a **while** loop. Remember, while loops loop until the condition is met. In this case the condition is `count < len(mol_code)`. In this example, the length of mol_code is 12, as benzene has 12 atoms.

But what is count here?

You'll see that each time round the loop we **increment** count by 1 using the `count = count + 1` instruction. `count` starts at zero and loops round going up by 1 until it is 11 (remember, python is zero-indexed, so `mol_code` of length 12 has elements from 0 to 11), when count is incremented the final time, the next time it goes round the loop `count = 12` and:

`12 < 12` is False so the loop stops. 

A brief aside, if you look up python code you will often see people using `+=` to mean, add 1 to this variable. i.e. 

`count += 1`

does the same thing as 

`count = count + 1`.

Most programmers prefer `+=` as it is fewer characters to type (and good programmers are lazy in a constructive manner!). But the in version above it can easy to mistake a `+=` for an `=` when you're scanning your code and this can make it harder to identify bugs.

Now that we've verified that our loop works correctly, we want to get the weights. 

```python
count = 0
while count < len(mol_code):
    atom = mol_code[count]
    weight = weight_dict[mol_code[count]]
    print(f'atom {count} is a {atom} has an atomic weight of {weight}')
    count = count + 1
```

Awesome, but we want the sum of those weights, so we add a new variable called mol_weight, and use it like count, but instead of adding `1` to it, we add the atomic weight we have read from the dictionary: 

```python
mol_weight = 0
count = 0
while count < len(mol_code):
    atom = mol_code[count]
    weight = weight_dict[mol_code[count]]
    print(f'atom {count} is a {atom} has an atomic weight of {weight}')
    count = count + 1
    mol_weight = mol_weight + weight
print(f'Molecular weight is {mol_weight}')
```

And now we've written this working script, we can wrap it in a function, so we can call it more easily:

```python
def calc_mol_weight(
    mol_code,
    weight_dict
    ):
    """
    Calculates molecular weight from a molecular code
    mol_code: molecular code - list of single character atoms
    weight_dict: input weight dict
    """
    mol_weight = 0
    count = 0
    while count < len(mol_code):
        atom = mol_code[count]
        weight = weight_dict[mol_code[count]]
        print(f'atom {count} is a {atom} has an atomic weight of {weight}')
        count = count + 1
        mol_weight = mol_weight + weight
    print(f'Molecular weight is {mol_weight}')
    
    return mol_weight
```

```python
mol_code = mol_dict['methanol']
calc_mol_weight(mol_code, weight_dict)
```

But, in fact, I'm going to make it even simpler, by reducing the number of inputs into that function, by putting the dictionary inside the function (after all, the mass of carbon is not going to change), this avoids me having to type out `weight_dict` each time and means that I can't break the function by changing `weight_dict` elsewhere in the file.

```python
def calc_mol_weight(mol_code):
    """
    Calculates molecular weight from a molecular code
    mol_code: molecular code - list of single character atoms
    """
    weight_dict = {'H': 1.008,
             'C': 12.011,
             'O': 15.999}
    mol_weight = 0
    count = 0
    while count < len(mol_code):
        atom = mol_code[count]
        weight = weight_dict[mol_code[count]]
        print(f'atom {count} is a {atom} has an atomic weight of {weight}')
        count = count + 1
        mol_weight = mol_weight + weight
    print(f'Molecular weight is {mol_weight}')
    
    return mol_weight
```

```python
mol_code1 = mol_dict['methanol']
calc_mol_weight(mol_code1)
```

```python
mol_code2 = mol_dict['ethanol']
calc_mol_weight(mol_code2)
```

Look back over the function above, this might have seemed impenetrable when you started this course, but now think how much you understand.