# 16. Numpy

While the basic python language is powerful, much of it's usefulness comes from the wide range of additional **libraries**. **Numpy** is one such python library (numpy being short for  *numerical python*) that adds in the basic numerical stuff that we need as scientists. 

To use numpy we must first import it. When we import it, most people rename it to `np`.

```python
# this imports numpy and gives it np as a nickname
import numpy as np
```

Now we can use numpy as we like. A good example is with the square root function that has been used previously.

```python
np.sqrt(121)
```

Remember lists? Lists were used a lot in the previous material, numpy uses a similar object called `array`. 

A list of values is a list, and can take any type of thing, numbers, words, objects. 

A numpy array of dimension 1 is a vector and is written as `np.array(my_list)` and can only take objects of the same type.

```python
my_list = [1, 2, 3, 4]
my_list
```

```python
my_1D_array = np.array([1,2,3,4])
my_1D_array
```

So far these are pretty similar. Where it starts to change is when you consider higher dimensional data sets (e.g. matrices).
You can build a list of lists, but it becomes a proper matrix if you convert it to a np.array, for example:

```python
my_list_of_lists = [[1, 2, 3], [4, 5, 6], [7, 8, 9], [10, 11, 12]]
my_list_of_lists
```

```python
my_2D_array=np.array(my_list_of_lists)
my_2D_array
```

You'll see one of the differences between the two data types here. The list datatype justs sees the length of the lists, rather than the shape of the overall array.

```python
len(my_list_of_lists)
```

```python
len(my_list_of_lists[0])
```

This is the shape of the matix in rows and columns. There are 4 rows and 3 columns:

```python
my_2D_array.shape
```

## Slicing

Rather than just **addressing** an array, we can **slice** the data, which was obsereved for lists. 

Using `:` grabs us all the data, for example:

```python
my_1D_array[:]
```

We can slice a range, for example, to grab the middle two numbers we would use `[2:4]` which would grab us from the 2nd to the 4th number

```python
my_1D_array[2:4]
```

We can slice from a point to the end of the array, for example, to get all the items from index 1 onwards, we'd write `my_1D_array[1:]`.

```python
my_1D_array[1:]
```

And to get all the items from the -1th (i.e. last) item backwards we'd do: `my_1D_array[:-1]`

```python
my_1D_array[:-1]
```

To remind you what `my_2D_array` is:

```python
my_2D_array
```

For arrays of higher than one dimension, such as `my_2D_array` we can slice the rows and columns separately. For example to get the column with index 1, we take all the rows `:` and column `1` as so:
`my_2D_array[:,1]`
(remembering that python indices go from 0 to 1):

```python
my_2D_array[:,1]
```

To get rows 2 onwards in column one, we'd write
`my_2D_array[2:,1]`

```python
my_2D_array[2:,1]
```

By combining various forms, we can select almost any subset of an array:

```python
my_2D_array[(1,3),1:]
```

## More on 2D arrays

Some operations can be done on the different dimensions of the array. In the example below, we use `np.sum()` which sums whatever it is given. In the first example it sums all the elements, in the second example it sums along the $0^{\mathrm{th}}$ (pronounced zeroth) axis, i.e., it sums along the rows to give you the sum for each column, and if we use `axis=1` it does the opposite, adding up the numbers along the columns to give you the sum for each row. If your array has more than 2 dimensions, you can do the same operations along any axis by using the appropriate axis number.


```python
# Compute sum of all elements; 
print(f"Sum all elements: np.sum(my_2D_array) gives {np.sum(my_2D_array)}")  
# Compute sum of each column; 
print(f"Sum all columns: np.sum(my_2D_array, axis=0) gives {np.sum(my_2D_array, axis=0)}")  
# Compute sum of each row; 
print(f"Sum all rows: np.sum(my_2D_array, axis=1) gives {np.sum(my_2D_array, axis=1)}")  
```

What's with this `axis=1` parameter?

It's not immediately obvious where axis comes from and why this needs to equal 1 for it to affect columns. To see why, just look at the `.shape` output:

```python
my_2D_array.shape
```

Remembering that Python starts counting from zero, then we have the first number as the number of rows (axis=0) and the second as the number of columns (axis=1).

## Other array operations

Multiplying arrays is done in a **element-wise** manner, i.e., the elements of each array are multiplied by each other (this is different from matrix multiplication).

```python
my_2D_array*my_2D_array
```

We can do matrix operations like transpose, like so:

```python
my_2D_array.T
```

This is how you would do matrix multiplication:

```python
np.matmul(my_2D_array, my_2D_array.T)
```

## Basic statistics in numpy

There are many other things that numpy can do, here's a few more examples. 
We can calculate means using `np.mean()`

```python
# this calculates the mean (average) of each column
means = np.mean(my_2D_array, axis=0)
means
```

and standard deviations as `np.std()`.

```python
# this calculates the standard deviation of each column
stds = np.std(my_2D_array, axis=0)
stds
```

And of course, it's easy to combine numpy operations to calculate things like the standard error, ste:
$$
ste = \frac{\sigma}{\sqrt{N}}
$$
where $\sigma$ is the standard deviation and $N$ is the number of samples. In the example below we calculate the standard error:

```python
# this calculates the standard error of each column
ste = np.std(my_2D_array, axis=0) / np.sqrt(my_2D_array.shape[0])
ste
```

You would use the standard deviation or standard error in error bar plots like the example below:
(Plotting is covered further in a later notebook).

```python
# import the matplotlib modules and nickname it plt
import matplotlib.pyplot as plt

# does an error bar plot, 
# first argument is the x values,
# second argument is the y values,
# yerr is the y error bars
# fmt is the style, - for line, o for circles, k for black
plt.errorbar([0,1,2], means, yerr=ste, fmt="-ok")

# labels x axis
plt.xlabel('Column no.')

# sets tick marks on x axis
plt.xticks([0,1,2])
# labels y axis
plt.ylabel('My data')
```

There are many many other things numpy offers, this is just a sample!

## Worked exercises 

Now your turn; in the exercises below, pick out the rows and columns and assign them to the correct variable, i.e. to assign column 1 to x you would enter `x = my_2D_array[:,1]`

```python
new_array=np.array([[1, 2, 3], [4, 5, 6], [7, 8, 9], [10, 11, 12], [13,14,15]])
new_array
```

## 1. Take the first column of data (column 0) and assign it to a variable called x1

So you should have `x1 = ...` where `...` is replaced by your code 

<!-- grade_id: cell-9b118de7be0df2c3, points: 0, grade: false, solution: true -->
```python
###
```

<!-- grade_id: cell-83fb30a7b954a8d4, points: 2, grade: true, solution: false -->
```python
assert (x1 == np.array([ 1,  4,  7, 10, 13])).all()
```

## 2. Take the 4th row of data (row 3) and assign it to a variable called x2

<!-- grade_id: cell-e60608fd68b6f3d0, points: 0, grade: false, solution: true -->
```python
###
```

<!-- grade_id: cell-0d4b03e174c35ce8, points: 2, grade: true, solution: false -->
```python
assert (x2 == np.array([10, 11, 12])).all()
```

## 3. Take the middle three numbers* and assign it to a variable called x3
*i.e. pull out the elements which contain the numbers 5, 8 and 11

<!-- grade_id: cell-49a5c8df89005586, points: 0, grade: false, solution: true -->
```python
###
```

<!-- grade_id: cell-735f589f66231f0c, points: 2, grade: true, solution: false -->
```python
assert (x3 == np.array([5,8,11])).all()
```

## 4. Calculate the square root of the sum of the elements in column 2 and assign it to a variable called x4

<!-- grade_id: cell-890a912eda399370, points: 0, grade: false, solution: true -->
```python
###
```

<!-- grade_id: cell-f2ba57c63763fb73, points: 2, grade: true, solution: false -->
```python
assert x4 - 6.708203932499369 < 0.00001
```

## 5. Calculate the means of the rows and assign it to a variable called x5

<!-- grade_id: cell-c08fbf2596db6e19, points: 0, grade: false, solution: true -->
```python
###
```

<!-- grade_id: cell-839d6bd112ed073b, points: 2, grade: true, solution: false -->
```python
assert (x5 == np.array([ 2.,  5.,  8., 11., 14.])).all()
```