# 20. RdKit 

Pandas is a powerful tool, and is very helpful for chemists.

**RdKit** is a computational chemistry module for Python and is extremely useful for chemists. 

In this notebook we are going to use Pandas and RdKit to find molecules with a certrain property in a dataset and display their structures.

First we import our modules

```python
# import Rdkit for chemistry:
import rdkit
# import Pandas for DataFrames:
import pandas as pd
# import matplotlib for plotting:
import matplotlib.pyplot as plt
# import NumPy for maths:
import numpy as np
```

Now we import our data for the *.csv* file:

```python
df = pd.read_csv("delaney.csv")
```

Now, lets look at our dataset using `head`, type the code to show the top of the dataset in the box below:

```python

```

You will see that this dataset is much more interesting than the previous ones you have looked at!

In the cell below type the code to find out how many molecules the dataset has. There are several different ways to do this, try using `len` or `tail`

```python

```

## SMILES and molecules as objects

If you remember in the dictionaries notebook we used a very simple method of inputting a chemical formula to allow us to use Python to calculate the molecular weight. There are nice methods or inputting chemical formulae into computers called things like SMILES, SMIRKS and SELFIES. These are strings which **encode** chemical formulae. Take a look at some of them in the cell below.

```python
df[['smiles']]
```

They might look like someone has leant on the keyboard, but you can learn to read them... (ish). The code below prints out the 5$^{\mathrm{th}}$ smiles string in the list:

```python
x=df['smiles'][4]
x
```

To read it:

there is a carbon ring, we know this as there is a 1 after the first carbon and a 1 at the end, indicating that the last carbon is linked to the first (*i.e.* the two carbons with 1 after them are joined). We know that there are 4 carbons and 1 sulfur in the ring. As they are lower case letters, we know that the ring is aromatic. Hydrogens are not included as it is assumed a chemist would be able to work out where they go from this information!

No one expects you to be able to read SMILES strings! What we can do is give the SMILES string to a computer program to read it. In this case we use a function from `rdkit`'s `Chem` function called `MolFromSmiles`, as below:

```python
rdkit.Chem.MolFromSmiles(x)
```

And we can see that this molecule is thiophene (and if you did not know that, you could have used the code below to read it from the entry in the DataFrame):

```python
df['Compound ID'][4]
```

This is also the first ever **object** that you have come across. Like real world objects it has **properties** and can do things, these are called **methods**. The output of `rdkit.Chem.MolFromSmiles(x)` is not a picture of a formula (although that is what is printed to the screen), it is an object. If you have ever played a computer game, you are familiar with objects, most objects you interact with in the game (blocks, guns, swords, spaceships, NPCs, angry birds, etc).

In the cell below you will define an object which you will call `thiophene_object`

```python
thiophene_object = rdkit.Chem.MolFromSmiles(x)
```

To see the structure, we print it to screen by putting it in a cell:

```python
thiophene_object
```

Now we use one of the built in **methods** called `GetNumAtoms`. As the method belongs to the object, we call it as `object.method()`, it is like a function, in that you can put specifications into the brackets, but you will not do that here:

```python
thiophene_object.GetNumAtoms()
```

This number does not seem right, trying another method, `GetNumHeavyAtoms`, counts the number of non hydrogen atoms:

```python
thiophene_object.GetNumHeavyAtoms()
```

So, in the `GetNumAtoms` the hydrogen atoms have not been counted:

```python
thiophene_object.GetNumAtoms() == thiophene_object.GetNumHeavyAtoms()
```


If you remember, the SMILES strings didn not have hydrogens in them. So we could use chemical knowledge to count the missing hydrogens...

or you can use `rdkit` to do it for you.

To do this you will use a rdkit method: `rdkit.Chem.AddHs` which will add the Hs to the structure. Call this method on the thiophene_object, by putting it in the `( )`, and then assigning this to a new object (`thiophene_object_H`): 

```python
thiophene_object_H = rdkit.Chem.AddHs(thiophene_object)
thiophene_object_H
```

Now the printed molecular structure has hydrogens on it. And you can check this by using similar code to above to count the number of atoms and the number of heavy atoms.

```python
thiophene_object_H.GetNumAtoms()
```

```python
thiophene_object_H.GetNumHeavyAtoms()
```

```python
thiophene_object_H.GetNumAtoms() == thiophene_object_H.GetNumHeavyAtoms()
```

It is now correct. You can count the number of hydrogens as well by doing some maths:

```python
thiophene_object_H.GetNumAtoms() - thiophene_object_H.GetNumHeavyAtoms()
```

## Investigating the dataset

Lets start by having a look at a few more molecules. 

You can put almost anything into a Python list, including objects. So, the code below makes a list of objects created from the first 12 molecules in the DataFrame. You do not need to understand the following **list comprehension** line of code, but it is worth looking at for a moment.

`df['smiles'][:12]` grabs the smiles column from the DataFrame, and just like in lists and arrays, we can use slices `:` to take parts of list. This takes the first 12 smiles strings (*note* that this logic can be applied in reverse with `df[:12]['smiles']`).

`rdkit.Chem.MolFromSmiles(x) for x` gives us the objects of the `MolFromSmiles` function called on each `x`. Where `x` is an element in the `df[:12]['smiles']` series. 

Finally as all of this code is in `[]`, this gives us a list of molecule objects.

```python
molecule_list = [rdkit.Chem.MolFromSmiles(x) for x in df['smiles'][:12] ]
molecule_list
```

The output of `molecule_list` is a list of objects (e.g. `<rdkit.Chem.rdchem.Mol at 0x2110b84a650>`) which tells us there is an rdkit Mol object and the weird looking number tells us exactly where on the computer it is stored (you will never need to know this!)

You can address this list as you can all other lists, for example here is the 9th object in the list:

```python
molecule_list[8]
```

But the really cool thing you can do with this is make a nice plot of the whole list, to do this you use `rdkit.Chem.Draw.MolsToGridImage` as below:

```python
rdkit.Chem.Draw.MolsToGridImage(molecule_list)
```

## Combining rdkit and Pandas

Lets continue investigating this dataset. 

To remind you, lets just print out the top of the DataFrame:

```python
df.head()
```

## Finding specific molecules

In the example below, you are going to find the molecule with the highest molecular weight and print its structure. 

Start by assigning the **'Molecular Weight'** column to a new variable called `this_column` (doing this makes it easier to copy code later).

```python
this_column = df['Molecular Weight']
this_column
```

Pandas has some useful built in methods like `idxmax` and `idxmin`, which will find the **i**n**d**e**x** of the **max**imum value in the column (*i.e.* the row number) and the **i**n**d**e**x** of the **min**imum value in the column.

So, call `idxmax` on `this_column`:

```python
max_idx = this_column.idxmax()
max_idx
```

We can use NumPy's `max` function to find out what the maximum **value** in this column is:

```python
np.max(this_column)
```

And then we can check that `idxmax` works as below:

```python
# goes to row 555 and reads out the Molecular weight column and shows the result in a DataFrame
df.iloc[[max_idx]][['Molecular Weight']]
```

And the entire row can be grabed by using:

```python
df.iloc[[max_idx]]
```

Finally, you can use the code below to pull out the SMILES string for the `max_idx` row defined previously and print the structure:

```python
x = df.iloc[max_idx]['smiles']
rdkit.Chem.MolFromSmiles(x)
```

## Worked problems

In the following questions, find the molecule that satisfies the constraints and plot its structure (as in the example above).

## Task 1: Polar surface area

Find which molecule has the maximum **polar surface area** and assign it to `max_idx_1` so you should have a solution like `max_idx_1 = ...` with the .. replaced with your code.

<!-- grade_id: cell-eaff66dcde612f40, points: 0, grade: false, solution: true -->
```python
###
```

<!-- grade_id: cell-4b057ca02b77c741, points: 3, grade: true, solution: false -->
```python
###
```

```python
# this prints out the information of molecule you found
df.iloc[[max_idx_1]]
```

In the cell below, write the code to print out the molecular struture of the molecule you have found.

<!-- grade_id: cell-b0d6b550b27d028d, points: 0, grade: false, solution: true -->
```python
###
```

## Task 2: Number of rotatable bonds

Find which molecule has the **maximum number of rotatable bonds** and assign it to `max_idx_2` so you should have a solution like `max_idx_2 = ...` with the .. replaced with your code.

<!-- grade_id: cell-96b4f92df87faf1a, points: 0, grade: false, solution: true -->
```python
###
```

<!-- grade_id: cell-ad9b45553657b506, points: 3, grade: true, solution: false -->
```python
assert max_idx_2
assert [x for x in df.iloc[[max_idx_2]]["Number of Rotatable Bonds"]][0] == 23
```

```python
# this prints out the information of molecule you found
df.iloc[[max_idx_2]]
```

In the cell below, write the code to print out the molecular struture of the molecule you have found.

<!-- grade_id: cell-8da654ab20680a67, points: 0, grade: false, solution: true -->
```python
###
```

## Task 3: Minimum rotatable bonds

Find which molecule has the **minimum number of rotatable bonds** and assign it to `min_idx_3` so you should have a solution like `min_idx_3 = ...` with the .. replaced with your code.

<!-- grade_id: cell-d5c04c9604d50943, points: 0, grade: false, solution: true -->
```python
###
```

<!-- grade_id: cell-dfde78dcb20a0b8a, points: 3, grade: true, solution: false -->
```python
assert min_idx_3
assert [x for x in df.iloc[[min_idx_3]]["Number of Rings"]][0] == 5
```

```python
# this prints out the information of molecule you found
df.iloc[[min_idx_3]]
```

In the cell below, write the code to print out the molecular struture of the molecule you have found.

<!-- grade_id: cell-d90c65c110e5e39e, points: 0, grade: false, solution: true -->
```python
###
```

## Task 4: No. of H-bond donors

Find which molecule has the **largest number of H-bond donors** and assign it to `max_idx_4` so you should have a solution like `max_idx_4 = ...` with the .. replaced with your code.

<!-- grade_id: cell-9d940d1422ac3686, points: 0, grade: false, solution: true -->
```python
###
```

<!-- grade_id: cell-fdc6e3b14e1f26a2, points: 3, grade: true, solution: false -->
```python
assert max_idx_4
assert [x for x in df.iloc[[max_idx_1]]["Number of Rings"]][0] == 3
```

```python
# this prints out the molecule you found
df.iloc[[max_idx_4]]
```

In the cell below, write the code to print out the molecular struture of the molecule you have found.

<!-- grade_id: cell-430315f2f8d05587, points: 0, grade: false, solution: true -->
```python
###
```

The answer from this question should tell you something about what polar surface area is and what causes it. Write down your obervations below.

```python
#
```

## Featurisation

The properties that we are looking at for each molecules can be called **descriptors** or **features**. Effectively this is a mathematical way of describing the molecule or prvoding a numerical way to code a feature of the molecule. The process of taking a physical object and making it machine readable is called **feautrisation**. Feautrisation could be as simple as taking experimental data for a molecule, however normally we want to featurise lots of molecules and relatively fast, therefore they are often obtained computationally. Also the features can range in complexity which can provide us with different levels of information. RDKit has several ways to featurise itself and this will be explored.

First of all we will import the descriptors module and run this for one molecule:

```python
from rdkit.Chem import Descriptors
```

```python
print(len(Descriptors._descList))
Descriptors._descList
```

```python
Descriptors.NumHDonors
```

Apply this to the `thiophene_object` does this match with the data from the `delaney.csv`?

```python

```

You can also obtain all the descriptors with `CalcMolDescriptors()` give it a try, what is the data type outputed? 

```python

```

```python
#
```

Write some code to extract the descriptors for each SMILES string in the delaney database and make a new dataframe of descriptor (`df_desc`). There are multiple ways to do this, you could use Pandas functionality or write a `for` loop, what if you wanted to also use the `Compound ID` in the final `df_desc`, could you enumerate to achieve this?

```python

```

```python
df_desc
```

Check the **qed** quantity of drug likeness - is it normally distributed? 

```python
#
```

There are other descriptors out there (*e.g.* Mordred), but you may have noticed that this takes a little while to calculate, because the descriptors are information rich, therefore for very large databases even these descriptors become computationally costly, and we may need simpler representations. Therefore it is possible to featurise a molecule just based solely on its structure. The simplest approach is to use a **Molecular Fingerprint**

A molecular fingerprint is a binary encoding of a molecule where the 0 and 1 represent some form of property. There are many different types of fingerprints and the details on how each one is constructed will not be discussed here. But RDKit does allow us to construct some molecular fingerprints, let's use the Morgan Fingerprint on the thiophene molecule. 

```python
from rdkit.Chem import rdFingerprintGenerator
```

```python
fpgen = rdFingerprintGenerator.GetMorganGenerator(radius=2, fpSize=2048)

thiophene_fp =  fpgen.GetFingerprintAsNumPy(thiophene_object)
thiophene_fp
```

We have a NumPy array (a vector) of 0 and 1's which provide a quick and easy coding for the molecules to represent its properties. Create a DataFrame called `df_fp` to get the fingerprints of every molecules in the `delaney.csv`

```python
df_fp
```

You will notice this is much quicker, but it is now more difficult to interpret the information in a chemically meaniningful way. There are also many types of fingerprints available (lots in RDKit with control on parameters). It is a good idea to investigate a few in projects!

These molecular fingerprints can be directly read into a ML algorithm. However once useful thing that can be done with them out the gate is to examine the "similarity" of molecules. The concept of similarity is a complex topic, one could view this is how close they are in **euclidean** "chemical space" with some of the RDKit descriptors. However, with binary data the **euclidean** distance does not have the same meaning. However mathematical forumla to quantiy similarity for this kind of data is commonly used. Again, we will not focus on the forumlation but a common approach is to use the **Tanimoto Similarity**. For example we will look at the similarity between Thiophene and Citral, first we now process the fingerprint as an object rather than a NumPy array. Perform a similar analysis to obtain `citral_fp_obj` and run the code to obtain the similarity. 

```python
thiophene_fp_obj =  fpgen.GetFingerprint(thiophene_object)
thiophene_fp_obj 
```

```python
#
```

```python
from rdkit import DataStructs

DataStructs.TanimotoSimilarity(thiophene_fp_obj,citral_fp_obj)
```

Are these objects similar? (note numbers range from 0 to 1). As a final activity find the molecule which is most similar to thiophene based on the Morgan fingerprint and Tanimoto similarity. Note you might want to make a new dataframe with thiophene removed as this will always come out as the most similar. You can use `df.drop()` to achieve this, and reset the index. Does this final result seem logical?

```python

```

```python

```

```python
#
```