# 3. Numbers and Maths in Python

Most programming languages can be used like fancy calculators.
It's possible to do all the arithmetic operations you learned about in school (and much, much more).

In this exercise, you will see how many of the standard arithmetic operators can be used in the Python programming language.

Let's list some of the operators:

* `+` plus
* `-` minus
* `/` slash (divide)
* `*` asterisk (multiply)
* `%` percent
* `<` less-than
* `>` greater-than
* `<=` less-than-or-equal-to
* `>=` greater-than-or-equal-to

*Do you notice any operator missing that you would expect to see in this list?*

To understand how these operators are used, let's run some more programs!

Hit `Shift+Enter` on each code block below to see the output from the calculation.

Note that, in a Jupyter Notebook the output of the last statement in a cell will automatically be printed for you.

```python
# Add 3 and 5
3 + 5
```

```python
# Subtract 8 from 10
10 - 8
```

```python
# Divide 100 by 4
100 / 4
```

In the above code blocks, notice that some of the operators are the same as those you have used since you first learned how to work with numbers. 

* `+` is used for addition
* `-` is used for subtraction

Other arithmetic operations are done using different symbols than you might be used to:

* `*` is used for multiplication instead of an 'x'
* `/` is used for division instead of the division symbol

Some symbols run arithmetic and logic operations that you might never have encountered before learning to code. See if you can make sense of the following operations by running the code blocks.

```python
# Check if 5 is less than 10
5 < 10
```

```python
# Check if 1200 is greater than or equal to 13
1200 >= 13
```

```python
# Check if 3 is greater than 89
3 > 89
```

Logic operators like **less-than** or **greater-than** elicit a `boolean` output value. \
That's a fancy way of saying you get an answer that is either `True` or `False`. 

See what happens when you combine some operators into more complex terms.

Run the code blocks below.

```python
3 + 2 > 2 + 3
```

```python
100*10 <= 1000
```

```python
30/3 > 1 + 2
```

You can also combine such arithmetic operations with the `print` statements that you used to write your first programs in **Exercises 1 and 2**.

```python
# Print a text string and the output of 4 multiplied by 4
print("The answer is: ", 4 * 4)
```

Use the code block below to write a comment line starting with a `#`, and `print` statement to output the answer of 40 divided by 10.

```python

```

These logic operators are not the only ones to output some (perhaps) unfamiliar results. See if you can understand what the `%` operator is doing by hitting `Shift+Enter` on the following code block.

```python
# Change this comment line to describe what the next line of code is doing.
10 % 3
```

```python
# Here are some more examples of using the '%' operator

print("The first example gives this answer: ", 200 % 11)
print("The second example gives this answer: ", 3 % 3)
```

### Exercise 3: Playing with the Program

It's time again to make up some more examples and explore the new code for yourself. Make some code. Break some code. Tinker with these arithmetic operators and try to understand their use (and misuse) in more detail before moving on. Here's some things to try:

1. What changes in the output when you change, say, `10 / 2` to `int(10 / 2)` ? Try both variations in the code block below and write a comment line to describe what you observed in the output.
2. Some operators can be 'repeated' to give more new functions. Try setting up two calculations, one each for demonstrating the use of `**` and `//`. Write a comment line above each example to describe the additional Python-enabled maths functions you have discovered.
3. For any of your calculations what happens if you take out a number and replace it with a letter by mistake?

Use the code blocks below to answer questions 1 and 2.

```python
### Question 1

```

```python
### Question 2

```

## Worked Problem

Below you will write the following line of code:

$x = 3^2 + 3\times2 $

This will then be checked with the use of an assert statement. **Assertions** are used when writing code to allow programmers to write **tests** to check that the code does what the programmer expects it to. Testing code is very important.

If the code is not correct then an error will be raised.

<!-- grade_id: cell-3e9e86d4ee4b3c9c, points: 0, grade: false, solution: true -->
```python
### Write your code
```

<!-- grade_id: cell-f092044555e2e69e, points: 1, grade: true, solution: false -->
```python
assert x > 0
assert x == 15
```