# 7. Logic

Life is full of interesting forks in the road. 

We can take turns this way `or` that, this way then that, this way `while` that, this way `if` that... all manner of twists and turns!

Far from the philisophical musings on life, we encounter such **'What If?'** moments whenever we play with the taps at a sink.

For example, `if` we turn the cold tap *then* cold water will run. That cold water will run `while` the tap remains open `and` will only stop `if` we  turn the cold tap back the other way.

Instead, `if` we turn on the cold tap `and` the hot tap, lukewarm water of a temperature determined by the flow rates of the hot `and` cold will run `while` the taps are open.

The deliberately highlighted words in the mundane example above serve to show just how we come across decisions, branch points, and **'what if?'** moments in all manner of everyday circumstances. 

In **notebook 8**, you will learn to appreciate the `boolean` data type more deeply. In simpler terms, the aforementioned forks in the road are determined by certain things proving `True` or `False`. 

**In this notebook, you will learn to determine when these `True` and `False` outputs will appear. Later, you will learn how to use these outputs to determine the path along the forks in the road your Python code will take.**

**Important: as you go through each cell below, think about what you think the result should be before hitting shift-enter to run the cell and reveal the answer. This will help you learn.**

### Play with the Program (Logic)

Use `Shift+Enter` to run each code block below.

The output will read `True` or `False`. 

The `not` keyword is applied in certain cases to essentially provide the opposite answer to whatever is computed in the brackets.

**Can you predict the output before you reveal it?**

```python
True
```

```python
False
```

```python
not True
```

```python
not False 
```

```python
True or False
```

```python
True or True
```

```python
False or True
```

```python
False or False
```

```python
True and True
```

```python
False and False
```

```python
not(True or False)
```

```python
not True or True
```

```python
not(True or True)
```

```python
not False or True
```

```python
not(False or True)
```

```python
not False or False
```

```python
not(False or False)
```

```python
not True and False
```

```python
not(True and False)
```

```python
not True and True
```

```python
not(True and True)
```

```python
not False and True 
```

```python
not(False and True)
```

### Equivalence
You can (and will) go far beyond the simple combinations of `True`, `False`, `and`, `or` and `not`, as you just did.

It is also possible to force `True` or `False` outputs from code that uses the symbols of **equivalence**.

**A quick reminder your use of `=` so far...**

In the likes of **notebook 5** you **assigned** various data objects to variables using the single `=` operator. This is not the same as the double `==` operator that you are about to use.

The `==` is like saying **equal to**.

Conversely, the related `!=` operator is used to state **not equal to**.

Let's try some examples and see how this works.

**Again, try to predict the output before pressing `Shift+Enter`**

```python
1 != 0
```

```python
1 != 1
```

```python
0 != 1
```

```python
0 != 0
```

```python
1 == 0
```

```python
1 == 1
```

```python
0 == 1
```

### Exercise 7: Play with the Program

In the first half of this notebook, you were subtly introduced to the difference in output that can occur when brackets appear in the logic statement. Let's try this out!

In the box below, assign **x** to the following logical statement:

`not (1 == 1 and 0 != 1)`

<!-- grade_id: cell-5a3f79ac913b1dd0, points: 0, grade: false, solution: true -->
```python
###
```

```python
x
```

In the box below, assign **y** to the following logical statement:

`not 1 == 1 and not 0 != 1`

<!-- grade_id: cell-351d0ddf6ee78b48, points: 0, grade: false, solution: true -->
```python
###
```

```python
y
```

<!-- grade_id: cell-a9c063a74503d1ec, points: 1, grade: true, solution: false -->
```python
assert type(x) == bool
x == False
```

<!-- grade_id: cell-ade9fa759821f373, points: 1, grade: true, solution: false -->
```python
assert type(y) == bool
y == False
```

Think about **why** you got the answers you did in this exercise. 

What is the effect of the brackets?

# Truth tables

## The maths and engineering of truth tables

Boolean logic underlies almost all of modern computing. It is a type of logic called **binary** logic, as in it has only 2 possible values and these are opposites to each other. Binary logic is binary logic regardless of the symbol used, you may well have seen this type of logic using 1 for True and 0 for False.

Here we are using Boolean logic in software to control the flow of the program. Its original use was in electronic engineering to control the flow of electrons, and it is still used that way in computer hardware today. For this reason, **truth tables** are built with inputs and output for the operation, as in electronic engineering these would literally be input wires to a **logic gate** which leads to output wire(s). 

There are 16 possible operations that can be done with binary logic that combines two inputs and gives one output, and 4 that can be done with binary logic and one input, one output. 

There are several sets of these logical operations that can be used to create all the others. The most common set is AND, OR, NOT. With these three operations (in maths) or logic gates (in computer hardware) you can perform any operation.

I will show you the truth tables for these operations below. To read them, think of:

Input 1 [logic] Input 2 $\longrightarrow$ Output

## AND operation

**Input 1 AND Input 2 $\longrightarrow$ Output**

The AND operation only gives a TRUE ouput if Input 1 and Input 2 are both TRUE. This matches what we expect from the English language, for example: 
Tomorrow is the 29th of the month is only TRUE if today is the 28th AND this month is not Februrary in a leap year.

The truth table for the AND operation is below.

| Input 1 | Input 2 | Output  |
|---------|---------|---------|
|    F    |    F    |    F    |
|    F    |    T    |    F    |
|    T    |    F    |    F    |
|    T    |    T    |    T    |

## OR operation

**Input 1 OR Input 2 $\longrightarrow$ Output**

The OR operation is a bit odd. It is true if Input 1 is true **or** Input 2 is true **or** Input 1 **and** Input 2 is true. It is used differently to how we use the word or in English. Technically, this is **inclusive OR**.

This is the truth table for the OR operation. 

| Input 1 | Input 2 | Output  |
|---------|---------|---------|
|    F    |    F    |    F    |
|    F    |    T    |    T    |
|    T    |    F    |    T    |
|    T    |    T    |    T    |

## NOT operation

**NOT Input $\longrightarrow$ Output**

This is a one input one output operation, and is simply the negative of the input. i.e.:

|  Input  |  Output |
|---------|---------|
|    T    |    F    |
|    F    |    T    |

The operations AND, OR and NOT are sufficient to recreate all the other Boolean operations and so are used in computer logic. 

## Other operations of note

### XOR operation

**Input 1 XOR Input 2 $\longrightarrow$ Output**

This is the **eXclusive-OR gate**, and this works how the word 'or' works in English. It is only True if input 1 is true or input 2 is true, but it is false if input 1 and 2 are true. 

| Input 1 | Input 2 | Output  |
|---------|---------|---------|
|    F    |    F    |    F    |
|    F    |    T    |    T    |
|    T    |    F    |    T    |
|    T    |    T    |    F    |

### NAND operation

**Input 1 NAND Input 2 $\longrightarrow$ Output**

What if we wanted to only use one operation and be able to recreate all the other Boolean operations used in computer logic, is it possible?

Yes, the NAND gate, negative AND, can do this. Here's the truth table. 

| Input 1 | Input 2 | Output  |
|---------|---------|---------|
|    F    |    F    |    T    |
|    F    |    T    |    T    |
|    T    |    F    |    T    |
|    T    |    T    |    F    |

We might want to use this logic if we only wanted to build a single type of gate. 

## Other logics

Binary is not the only logic we can use. Here are a few other examples:

1. trinary: True, False, Neither True nor False
2. quaternary: True, False, Neither True nor False, Both True and False
3. fuzzy: uses numbers rather than binary, i.e. 80% true, 20% false

Binary comes from the Western philosophical tradition (think Ancient Greek philosophers) and quaternery logic comes from the Eastern philosophical tradition. You might want to think about how trinary or quaternary might be used.


# Truth tables in chemistry

Ok, so this is rather abstract, does it have any relation to chemistry? Actually, yes!

Modern computers are built from materials, and those materials have been developed and designed by chemists. This requires finding materials that can 'switch' in some way, and all modern computing is based on changing the amount of electrons flowing through a material. 

This is an exciting, current (pun-intended) area of chemistry and materials research.

Some examples:

1. Attempts to shrink transistors (the go-to device for doing computer logic in hardware) to allow for smaller, faster, more-powerful, cheaper and greener computers. Shrinking their size involves a lot of nanotechnology and new materials
2. CCDs, these are the components found in all digital cameras, originally designed to replace transistors, instead their reaction to light meant that the technology was used in cameras instead
2. Spin-tronics: making devices that switch based on the spin of electrons, these are currently research devices but have uses in hard discs and data storage
3. Memristors and ReRAM: devices that implement one of the other Boolean logic operations called implication, IMP. These are mostly research devices at the moment, but they have a huge potential in make greener computers and computers that operate more like we do
4. Quantum computers: using quantum states to implement logic
5. Molecular electronics (and some molecular machines): using molecules to implement logic or switching

Chemistry is essential to research in these areas.