# 9. Defining Functions in Python

Every piece of code you have written in this course has shared one major point in common: all the code is written linearly, and is executed from the first line, in order, all the way to the very last line of code.

That changes here.

Now, you learn about **writing your own functions** in Python. Before we go there, let's look at some familiar code through a new lens. Hit `Shift+Enter` on the code block below.

```python
print("Hello world!")
```

In the beginning, we never stopped to pick this code apart. We were too busy with the excitement of writing our very first program.

Now, we will pay a little more attention to how the above line of code is constructed.

* The name of the function is `print`

* We signal its execution (we click 'go!') by the use of brackets `()`.

* The `print()` function, like many others, accepts **arguments** - extra stuff inside the brackets that gets fed into the function to tune what it does.

* The `print` function also works *without* any extra arguments. However, we still need the `()` to signal to the interpreter that we want the result of *calling* the function and not just the function itself.

Hit `Shift+Enter` on the following code blocks.

```python
# This print command has no input arguments between the brackets. It prints a blank line.
print()
```

```python
print("This line of text is separated from another line of text by two blank lines, each blank line printed by a different technique.")

print("\n")

print()

print("Here's the second line of text!")

# if we don't include () we just get the function.
print
```

A function is simply a smaller python program that you write and then reference. Let's create a simple function **definition** that will print two things and then **call** it ourselves.

Hit `Shift+Enter` on the code blocks below.

```python
# This function is called print_two and takes two arguments
def print_two(a, b):
    """ prints two arguments
    a: first argument
    b: second argument
    """
    print(f"First argument in print_two function definition is: {a}")
    print(f"Second argument in print_two function definition is: {b}")
    print("\n")
    return

# Call the functions
print_two("Hello again, ","world!")
```

```python
# This function, print_none, takes no arguments
def print_none():
    """a broken print function!"""
    print("I got nothin.")
    return

print_none()
```

*Most* python functions, and all the ones you will write in this course, take a fixed number of arguments* - if you pass the wrong number, then the code will throw an error. For example, if you wrote:

`print_none("Please, please print me (print me!)")`

You will get the error :

`TypeError: print_none() takes 0 positional arguments but 1 was given`

which tells you that the function `print_none()` takes no arguments, but you gave it one argument (the string `"Please, please print me (print me!)"`), so this is an error as you have the wrong type of input. 

*The astute among you will realise that the built-in `print` function is one of the special functions that can take multiple numbers of arguments.

Print functions like we have looked at so far are actually a special type of function that do not return a value. These are sometimes called *procedures*.
Most functions act like mathematical functions and do something to their arguments before *returning* the resulting value. Consider the following function to add one to the number you've passed it:

```python
def add_one(input_value):
    """adds one to input value and returns that
    """
    return input_value + 1

x = 5
print(f"x is {x}")
# Note that you can safely store the result of a function in one of the passed in variables!
x = add_one(x)
print(f"x is now {x}")
# Also note that the print function can take any valid python statement as an inline parameter
print(f"and one more than that is {add_one(x)}")
```

Whilst some comments are provided in the code block above, **do not worry if this doesn't make any sense yet**! Here are some things to note right away:

* We use `def` to **define** a function.
* The name of the function (whatever you like) follows the `def` keyword.
* Brackets follow the function name.
* If your function needs some input to work properly, these can be included as **arguments** inside the brackets.
* All opening lines of a function definition (i.e. the line that has `def` on it) ends with a colon `:`.
* **Everything inside** a function is **indented** in order to signify what is in the function defintion and what is not.
* You *may* use block comments (""") to describe what your function does immediately after the definition. You do not have to do this, but should for complex or important functions, or ones you expect other people to use - these are called **docstrings**.
* If a function will **return** something, you use the *return* keyword with the value to return. This value comes out of the function so you can use it. For example, in the code above, x was set to 6 after calling the function, since the function returned the value 6 (5+1) and x was assigned to the output of the function

### Why bother with functions?

Throughout this introductory course, you have been asked to fix some bugs in your code.

The odd typo here, the missing bracket there... once you get to grips with the language, bugs are relatively easy to spot in code that's less than, say, 20 lines long.

Manually scanning for code bugs *will not* work when you get to more complex scientific code that can be thousands, or even *millions* of lines long!

Functions enable you to create little 'bubbles' of self-contained code. You can pass stuff in and get new stuff out.

The useful thing is that, if your program has a bug, functions make it much easier to track the bug and fix it!

The job of rewriting code to introduce functions rather than linear code is known as *refactoring*

### A real computer code bug!

It's easy to use the term 'bug' as it relates to computer code without ever questioning *why* such code errors are called bugs. There is a reason!

In 1947, a computer at Harvard University malfunctioned because a moth (a real life bug) got stuck inside.

<img src=https://substackcdn.com/image/fetch/$s_!oZMv!,w_1456,c_limit,f_webp,q_auto:good,fl_progressive:steep/https%3A%2F%2Fbucketeer-e05bbc84-baa3-437e-9518-adb32be77984.s3.amazonaws.com%2Fpublic%2Fimages%2F0660dd6a-54f9-41f8-bf54-beaf3f47c532_698x580.jpeg>

*Photograph credit: NAVAL SURFACE WARFARE CENTER, DAHLGREN, VIRGINIA*

Now back to learning about functions. Here's the same exemplar code as shown earlier in this exercise. Two questions:


1. What lines are missing in the following code compared to the code shown earlier in the exercise? Write your answer as a comment line in the empty code block below.

2. What effect do these missing lines have on the output compared to the otherwise identical code block above? Hit `Shift+Enter` on the block below...

```python
# This function is called print_two_again and takes two arguments
def print_two_again(a, b):
    """prints two arguments
    a: first argument
    b: second argument
    """
    print(f"First argument in print_two function definition is: {a}")
    print(f"Second argument in print_two function definition is: {b}")
    print("\n")

```

write your answer here:



```python
#
```

Spoiler alert - a whole lot of nothing! There is nothing printed when you run this code block.

But we still have `print` statements in this code...so why is nothing printing from this code block?

Hmmm...

Here is the most important new point to learn about functions, as demonstrated by the lack of any output from the last code block:

The functions have been **defined**... but **not called**!

Look again at the lines that were removed from the second version of the code above.

### Exercise 9: Play with the Program

It is definitely not expected that you will have mastered `def` and functions in a single notebook.

In this exercise, we look at some more functions and how you can use them to write more powerful programs.

There is one main learning point from this exercise before we get to any code:

***Variables defined inside a function are not connected and not the same as variables running outside the function, in the main part of your program.***

Run the code block below with `Shift+Enter` so we can start to understand the difference between variables **inside** and **outside** a function. Read the comment lines to help with your understanding.

```python
# Define a function that takes 2 input arguments
# Everything after the colon that is INDENTED is part of the function.
def grams_per_litre(grams,litres): 
    """Calculates concentration"""
    
    # Define a local variable that holds the result of a calculation used to find the grams per litre value.
    # Round the answer to 3 decimal places.
    # Print the answer and format the output so that the number displays only 2 decimal places.
    
    conc = round((grams / litres),3)
    print(f"The concentration in grams per litre is: {conc}", "g/L")
    return conc

# define two new variables in the main script (outside the function) that hold number of grams and number of litres.    
grams = 50
litres = 3

# call the function and pass in the two variables above as input arguments to run the function properly.
grams_per_litre(grams, litres)
```

In the above code block, you have performed three main tasks:

1. Defined a function called `grams_per_litre` that takes two input arguments inside the brackets.
2. Defined some variables in the main script. These variables were named `grams` and `litres`.
3. Called the `grams_per_litre` function, passing in the two variables defined in the main script to act as the two necessary  input arguments for this function.

Here's the exact same code in the block below, with one minor but instructive change. Can you see what has been changed? Write your observation as a comment line in the blank code block below directly below this one:

```python
# Define a function that takes 2 input arguments
# Everything after the colon that is INDENTED will be executed as part of the function (or 'mini script').
def grams_per_litre(grams,litres): 
    """Calculates concentration"""
    
    # Define a local variable that holds the result of a calculation used to find the grams per litre value.
    # Round the answer to 3 decimal places.
    # Print the answer and format the output so that the number displays only 2 decimal places.
    
    conc = round((grams / litres),3)
    print(f"The concentration in grams per litre is: {conc}", "g/L")
    return conc

# define two new variables in the main script (outside the function) that hold number of grams and number of litres.    
x = 50
y = 3

# call the function and pass in the two variables above as input arguments to run the function properly.
grams_per_litre(x, y)

```

write your answer here:


```python
#
```

In the first code block, the names of the variables outside the function had the same names as the input arguments and variables defined in the function. 

`grams` and `litres` were used to hold the integers `50` and `3`...

...but then those same names, `grams` and `litres` were used to define the `grams_per_litre` function.

In the ***second*** code block, the input arguments and variables inside the `grams_per_litre` function are still named `grams` and `litres`. 

However, the separate variables defined outside the function are now named `x` and `y`. When the `grams_per_litre` function is **called**, the values held in `x` and `y` are **passed into the function** by including these variables inside the brackets.

**The output from both variations of the code in the two blocks above is EXACTLY THE SAME!**

The key point here is that, just because the function definition, specificed with `def`, uses certain names for the arguments in the brackets, that does not mean that these names need to be used when the function is called.

Although you can, it's generally a good idea to not reuse variable names inside functions (outside of loop variables - we'll explain about these later). The reason for this is it becomes easy to forget which version of a variable you are referring to in your code.


Use the code block below to revise the difference between **defining** a function and **calling (i.e. running)** a function.

Complete the comment lines in your own words next to each `#` in the code below.

Be clear about where the `cheese_and_crackers` function is **defined** and where it is **called**.

*(Hint - the function has been called more than once. There are also some lines that are neither about defining or calling the function. Add comments to these other lines, too).*

```python
# 
def cheese_and_crackers(cheese_count, box_of_crackers): 
    print(f"You have {cheese_count} cheeses.")
    print(f"You have {box_of_crackers} boxes of crackers.")
    print()
    return


# 
cheese_and_crackers(20,30)

#
amount_of_cheese = 10
amount_of_crackers = 20

# 
cheese_and_crackers(amount_of_cheese,amount_of_crackers)


# 
cheese_and_crackers(10+20,5+6)


# 
cheese_and_crackers(amount_of_cheese+10,amount_of_crackers+1000)
```

### Worked Problem 1: 

Complete the function below to create a function to add two numbers together and return the sum, i.e.:

$y = a + b$ 

<!-- grade_id: cell-e5bc357ea0059391, points: 0, grade: false, solution: true -->
```python
# This function is called add_two and takes two arguments
def add_two(a, b):
    """adds two arguments
    a: first argument
    b: second argument
    """

    ###
    
    return y
```

```python
# you can test your function here
add_two(10,1)
```

<!-- grade_id: cell-e4732c2365719edb, points: 1, grade: true, solution: false -->
```python
# these assert what the solution should be to check your function
assert add_two(a=1, b=1) == 2
assert add_two(a=-1, b=1) == 0
```

### Worked Problem 2: 

Write a function to concatinate (i.e. stick) two strings together.

For example, if you were to input: 
'Hello' and 'world'

the function should return: 
'Hello world'

You will have to put the space in yourself. The string version of a space is a space in hyphens, i.e. `' '` is a space.

<!-- grade_id: cell-c16b92c37ea9e9d4, points: 0, grade: false, solution: true -->
```python
# This function is called print_two and takes two arguments
def add_two_strings(a,b):
    """prints two arguments
    a: first string
    b: second string
    """

    ###
    
    return y
```

```python
# you can test your function here
add_two_strings('Hello','world')
```

<!-- grade_id: cell-71302cb095acf789, points: 1, grade: true, solution: false -->
```python
assert type(add_two_strings('hello', 'world')) == str
assert add_two_strings('hello', 'world') == 'hello world' or add_two_strings('hello', 'world') == 'helloworld'
```

<!-- grade_id: cell-91ad1e983f616b05, points: 1, grade: true, solution: false -->
```python
assert add_two_strings('hello', 'world') == 'hello world'
assert not add_two_strings('egg', 'nog') == 'hello world'
assert add_two_strings('egg', 'nog') == 'egg nog'
```

### Worked Problem 3:

Write a function to calculate the hypotenuse using Pythagorean theorem:

$$c = \sqrt(a^2 + b^2)$$

To do this you will need to use a function to calculate the square root of a number. We need to import this function from a package called `numpy` (numerical python), and this package is imported and given the shortname `np` by calling `import numpy as np`. You only need to do this once and the standard convention is to import all the packages we need at the top of file.

To calculate the square root of a number, `x`, you would type, `np.sqrt(x)`, e.g. `np.sqrt(4)` would give you 2. Use the cell below to try out numpy's square root function.

```python
import numpy as np
```

```python
np.sqrt(4)
```

**Finish the function below so that it calculates the hypotenuse**

<!-- grade_id: cell-d47cc09a584d2d8f, points: 0, grade: false, solution: true -->
```python
# This function calculates the hypotenuse
def calc_hypotenuse(a,b):
    """uses pythagorus
    a: first argument
    b: second argument
    """

    ###
    
    return c
```

```python
calc_hypotenuse(3,4)
```

<!-- grade_id: cell-b0604bbfe576ea75, points: 3, grade: true, solution: false -->
```python
assert calc_hypotenuse(3,4) == 5
assert calc_hypotenuse(9,3) - 9.486832980505138 < 0.00001
```

### Worked Problem 4

In this section we are going to write a function that takes in a temperature in kelvin and returns the temperature in celsius. 

**Fill in the missing parts of the function below**

<!-- grade_id: cell-0a2ec36dc274d9ca, points: 0, grade: false, solution: true -->
```python
def kelvin_to_celsius(kelvin):

    ###

# define an input temperature in kelvin
absolute_temp = 298.15

# use the function to convert this to Celsius
celsius_temp = kelvin_to_celsius(absolute_temp)

print(f"The temperature in Celsius is: {celsius_temp}")
```

```python
# you can try your function out here
kelvin_to_celsius(50)
```

<!-- grade_id: cell-779572c526b3c8f3, points: 2, grade: true, solution: false -->
```python
assert celsius_temp == 25
assert kelvin_to_celsius(4) == -269.15
```

### Worked Problem 5

Taking inspiration from what you have just done, in this section write a function that takes in a temperature in celsius and returns the temperature in kelvin.

**Fill in the missing parts of the function below and run the function.**

**call the function `celsius_to_kelvin`**

**run celcius_to_kelvin with `celsius_temp` as the input (with `celsius_temp` set to -78) and set the output to `absolute_temp`**

<!-- grade_id: cell-bb25194e452d57d2, points: 0, grade: false, solution: true -->
```python
###

print(f"The temperature in Kelvin is: {absolute_temp}")
```

```python
# you can try your function out here 
celsius_to_kelvin(4)
```

<!-- grade_id: cell-0419b6f0f57f7b13, points: 2, grade: true, solution: false -->
```python
assert celsius_temp == -78
assert absolute_temp - 195.15 < 0.01
assert celsius_to_kelvin(4) ==  277.15
```

**Excellent! You have just completed another pivotal lesson in learning Python. Defining functions and understanding their structure is a huge step.**